# PS06

DS256, Gettysburg College

Prof Eatai Roth

Due: Nov 3, 2025 5p

**Your Name**: ...

**Collaborators**: ...

By submitting this work, you attest that: 
 - it reflects your own effort and understanding of the material
 - you did not ask for help from students outside the course, other than PLAs
 - you did not use a generative AI to produce any code used below.

**1.**

Looking again at the dataframe of election results. It's not formatted in the most convenient way. Let's improve the organization of this dataframe.

Goal:
 - One row for each county-and-year combo.
 - Within a row, there should be columns for each party candidate and the number of votes they received.
 - No duplicated or unnecessary data.



```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt

```

```python

elec_df = pd.read_csv('https://raw.githubusercontent.com/GettysburgDataScience/datasets/refs/heads/main/us_political/countypres_2000-2024.csv')

elec_df.head()
```

Output:
```
year    state state_po county_name  county_fips        office  \
0  2000  ALABAMA       AL     AUTAUGA       1001.0  US PRESIDENT   
1  2000  ALABAMA       AL     AUTAUGA       1001.0  US PRESIDENT   
2  2000  ALABAMA       AL     AUTAUGA       1001.0  US PRESIDENT   
3  2000  ALABAMA       AL     AUTAUGA       1001.0  US PRESIDENT   
4  2000  ALABAMA       AL     BALDWIN       1003.0  US PRESIDENT   

        candidate       party  candidatevotes  totalvotes   version   mode  
0         AL GORE    DEMOCRAT            4942       17208  20250821  TOTAL  
1  GEORGE W. BUSH  REPUBLICAN           11993       17208  20250821  TOTAL  
2           OTHER       OTHER             113       17208  20250821  TOTAL  
3     RALPH NADER       GREEN             160       17208  20250821  TOTAL  
4         AL GORE    DEMOCRAT           13997       56480  20250821  TOTAL
```

**2.** 
The remainder of the assignment uses the dataframe you re-organized above in conjunction with two other dataframes:
 - ```ev_df``` - the Washington state electric vehicle dataframe from last class
 - ```agesex_df``` - a dataframe of age and sex demographics for all US counties.

```python
ev_url = 'https://raw.githubusercontent.com/GettysburgDataScience/datasets/refs/heads/main/Electric_Vehicle_WA.csv'
ev_df = pd.read_csv(ev_url)

agesex_df = pd.read_csv('https://raw.githubusercontent.com/GettysburgDataScience/datasets/refs/heads/main/us_political/stats_america/Population-by-Age-and-Sex/Population%20by%20Age%20and%20Sex%20-%20US%2C%20States%2C%20Counties.csv')

display(ev_df.head())
display(agesex_df.head())
```

Output:
```
VIN (1-10)     County     City State  Postal Code  Model Year    Make  \
0  WA1E2AFY8R   Thurston  Olympia    WA      98512.0        2024    AUDI   
1  WAUUPBFF4J     Yakima   Wapato    WA      98951.0        2018    AUDI   
2  1N4AZ0CP0F       King  Seattle    WA      98125.0        2015  NISSAN   
3  WA1VAAGE5K       King     Kent    WA      98031.0        2019    AUDI   
4  7SAXCAE57N  Snohomish  Bothell    WA      98021.0        2022   TESLA   

     Model                   Electric Vehicle Type  \
0     Q5 E  Plug-in Hybrid Electric Vehicle (PHEV)   
1       A3  Plug-in Hybrid Electric Vehicle (PHEV)   
2     LEAF          Battery Electric Vehicle (BEV)   
3   E-TRON          Battery Electric Vehicle (BEV)   
4  MODEL X          Battery Electric Vehicle (BEV)   

   Clean Alternative Fuel Vehicle (CAFV) Eligibility  Electric Range  \
0              Not eligible due to low battery range            23.0   
1              Not eligible due to low battery range            16.0   
2    
```

**2**
    
 - a. What is the fips of the PA county with the greatest republican lean in 2024?
 - b. In the US in 2010, what is the fips (same as IBRC_Geo_ID) of the county that had the highest percentage of individuals aged 65+?
 - c. In Washington state 2024, how many electric vehicles were in the county with the highest percentage of people in the age range 18-44?
 - d. Rounded to the nearest percentage, what was the lean of that county (as a positive number)?
 - e. Find the most popular make of electric car in the Washington state county that has the greatest democratic lean. How many vehicles of that make are registered in that state?
 - f. Create three scatter plots, for each of 2016, 2020, 2024 in which each data point represents a county, the x-axis is percentage of people age 18-24 and the y-axis is political leaning of the county.

Tests for q1 (points: None)
```
>>> entree = {'Burger': 5, 'Taco': 3}
>>> side = {'Fries': 2, 'Salad': 1}
>>> assert combo_menu(entree, side) == {'Burger and Fries': 7, 'Burger and Salad': 6, 'Taco and Fries': 5, 'Taco and Salad': 4}
>>> assert combo_menu({}, {'Fries': 2, 'Salad': 1}) == {}
>>> assert combo_menu({'Burger': 5, 'Taco': 3}, {}) == {}
>>> assert combo_menu({'Pasta': 8}, {'Bread': 2}) == {'Pasta and Bread': 10}
>>> entree = {'Promo': 0, 'Discount': -1}
>>> side = {'Small': 0, 'Freebie': 0}
>>> assert combo_menu(entree, side) == {'Promo and Small': 0, 'Promo and Freebie': 0, 'Discount and Small': -1, 'Discount and Freebie': -1}
```


Tests for q2 (points: None)
```
>>> inv = {'apple': 5, 'orange': 2}
>>> res = update_inventory(inv.copy(), 'banana', 3)
>>> assert res == {'apple': 5, 'orange': 2, 'banana': 3}
>>> inv = {'apple': 5, 'banana': 3}
>>> res = update_inventory(inv.copy(), 'apple', 2)
>>> assert res == {'apple': 7, 'banana': 3}
>>> inv = {'apple': 5, 'banana': 3}
>>> res = update_inventory(inv.copy(), 'apple', -2)
>>> assert res == {'apple': 3, 'banana': 3}
>>> res = update_inventory({'apple': 3}, 'apple', -3)
>>> assert res == {}
>>> res = update_inventory({'apple': 3}, 'apple', -10)
>>> assert res == {}
>>> res = update_inventory({}, 'pear', 0)
>>> assert res == {}
>>> res = update_inventory({}, 'pear', -1)
>>> assert res == {}
>>> inv = {'a': 1}
>>> res = update_inventory(inv, 'b', 2)
>>> res = update_inventory(inv, 'a', -1)
>>> assert res == {'b': 2}
```


Tests for q3 (points: None)
```
>>> roster = {'Argyle': {'ps01': 9, 'ps02': 10, 'ps03': 'M'}, 'Belinda': {'ps01': 10, 'ps02': 8, 'ps03': 'M'}, 'Camilla': {'ps01': 10, 'ps02': 10, 'ps03': 'M'}, 'Dana': {'ps01': 'M', 'ps02': 5, 'ps03': 'M'}}
>>> assert abs(assignment_mean(roster, 'ps01') - (9 + 10 + 10) / 3) < 1e-08
>>> assert assignment_mean(roster, 'ps02') == 8.25
>>> assert assignment_mean(roster, 'ps03') == 'M'
>>> assert assignment_mean(roster, 'ps04') == 'M'
```
